# COGS 108 - Assignment 1: Setting Up

This assignment covers Python, git, and Github. 

Visit the relevant tutorials on https://github.com/COGS108/Tutorials for guidance for these topics.

This assignment is out of 10 points, worth 10% of your grade.

## IMPORTANT

### In the markdown cell below, write in your Student ID and GitHub username. 

The details you enter below will be how we get your account details to check that you have completed the GitHub portions of this assignment.

Please do not remove anything from the cell below, just add in the requested information to each line.

My Student ID number is: A13561164

My Github username is: xih108

^Make sure this is your username, **not the email address** you used to create the account.

The following couple cells are checks that you are running the right tools for these assignments.

If either of them fail, you will need to update some things. If you are not sure how, ask in office hours.

```python
# These assignments require Python 3 (not Python 2).
# Run this cell to check your version of Python.
import sys
assert sys.version_info[0] == 3, "You must use Python 3 (preferably 3.6 or 3.7) for these assignments."
```

```python
# Run this cell to check your version of IPython. 
# If you get an error (the message prints out), stop and upgrade Jupyter.
import IPython
assert IPython.version_info[0] >= 3, "Your version of IPython is too old, please update it."
```

## Part 1 (5 points)

Follow along on the instructions [here](https://github.com/COGS108/Overview/blob/master/A1_Setup.pdf) to do Part 1 of the assignment with git/GitHub.

## Part 2 (5 points)

This part of the assignment is focused on some practice with Python, and with practicing working with the format of the assignments. 

This class assumes some prior knowledge of Python. In the following questions, you will need to work with basic (standard library) data types (floats, lists, dictionaries, etc.) and control flow (conditionals, loops, functions, etc). If the questions in this section are totally unfamiliar to you, you may need revisit some practice materials to catch up with some of the programming. 

Through these questions, we will also prompt you to use a couple slightly more advanced standard library functions (for example, 'enumerate' and 'zip'), that may be new to you.

Each question should be answerable with a relatively small number of lines of code, up to about 5-7 lines at most. 

If you are having any trouble, remember to visit the course tutorials (https://github.com/COGS108/Tutorials). Assignment questions often follow the structure of examples provided in the tutorials, and a large number of relevant links and external materials are also indexed in the tutorials. 

### How to complete assignments

Whenever you see:

```python
YOUR CODE HERE
raise NotImplementedError()
```

You need to replace this section with some code that answers the questions and meets the specified criteria. Make sure you remove the 'raise' line when you do this (or your notebook will raise an error, regardless of any other code, and thus fail the grading tests).

Recall that any cell with one or more `assert` statements in it is a test cell. You should not try to change or delete these cells.

Note that there might be more than one `assert` for each question. 

If a test does fail, read the error that is printed out. The error should let you know which test failed, which may be useful for debugging.

```python
# PRINTING VARIABLES
# A reminder that you can (and should) print and check variables as you go.
#  This allows you to check what values they hold, and debug if anything unexpected happens.

# Define a variable
math_result = 2 * 4

# Print out the value(s) of a variable.
print(math_result)
```

Output:
```
8
```

**Question 1: Defining variables (0.5 points)**

Define a tuple called `var_a`, that contains the numbers 1-10 (inclusively).  
Define a list called `var_b`, that contains individual letters a-j (inclusively).


<!-- grade_id: cell-4db8c933a259ed1b, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
var_a = (1,2,3,4,5,6,7,8,9,10)
var_b = ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']
```

<!-- grade_id: cell-88582fbb27969b63, points: 0.5, grade: true, solution: false -->
```python
# Tests for Q1

# These tests check the variables are defined
assert var_a
assert var_b

# These tests check that the variables are the right data types
assert isinstance(var_a, tuple)
assert isinstance(var_b, list)

# These tests check the variables have the right values
assert var_a == (1, 2, 3, 4, 5, 6, 7, 8, 9, 10)
assert var_b == ['a', 'b', 'c', 'd', 'e', 'f', 'g', 'h', 'i', 'j']
```

**Question 2: Defining Variables, Part II (0.5 points)**

Create a Python dict called `dictionary` where the keys are the elements in `var_b` and the values are the corresponding elements in `var_a`.

The `zip` function may be useful.

You might also make use of a Python dictionary comprehension.

<!-- grade_id: cell-a0c9ee8ea65481d3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
dictionary = dict(zip(var_b,var_a))
```

<!-- grade_id: cell-611de86b26434b5e, points: 0.5, grade: true, solution: false -->
```python
# Tests for Q2
assert dictionary

# There are also some hidden tests that check the contents of dictionary
```

**Question 3: Control Flow (1 point)**

Loop through the provided list `lst`. For each element, check if it is an even number. If the element is an even number, append the INDEX of that element to the list `inds`.

Note that you are adding the index to `inds`, **not the element itself**.
   
Hint: to check if a number is even, you can use the modulo `%` operator.

Hint: to loop through an iterable, keeping track of the index, you can use the `enumerate` function.

<!-- grade_id: cell-5c803671612c731a, points: 0, grade: false, solution: true -->
```python
# These variables are provided to you.
lst = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15]
inds = []

# YOUR CODE HERE
inds = [i for i in range(len(lst)) if lst[i] % 2 ==0]
```

<!-- grade_id: cell-90cd314531614439, points: 1, grade: true, solution: false -->
```python
# Test for Q3
assert inds

# There is also a hidden test that checks the contents of 'inds'
```

**Question 4: Functions (1 point)**

Write a function `squared_diff` that takes two number inputs and returns the squared difference of the two numbers i.e. $(a - b)^2$. For example:

```python
>>> squared_diff(2, 4)
4
>>> squared_diff(10, 1)
81
```

<!-- grade_id: add2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
def squared_diff(a,b):
    return (a-b)**2
```

<!-- grade_id: add2-tests, points: 1, grade: true, solution: false -->
```python
# Tests for Q4
assert squared_diff(2, 4) == 4
assert squared_diff(10, 1) == 81

# There are also hidden tests that checks more cases of squared_diff
```

**Question 5: Putting it all together (1 point)**

Create a dictionary named `second_dictionary`. It should have the key-value pairs as `dictionary`, *unless the original value is odd*. If so, the value in `second_dictionary` should be the squared difference of the original value and `10`.

Note: to loop through key-value pairs in a dictionary, check out the `.items` method.

<!-- grade_id: cell-50b7166b6554cfff, points: 0, grade: false, solution: true -->
```python
# This variable is provided to you.
second_dictionary = dict()

# YOUR CODE HERE
for i in dictionary:
    if dictionary[i] % 2 == 1:
        second_dictionary[i] = squared_diff(dictionary[i],10)
    else:
        second_dictionary[i] = dictionary[i] 
```

<!-- grade_id: cell-234cdbef12134b0d, points: 1, grade: true, solution: false -->
```python
# Tests for Q5

assert second_dictionary

# There are hidden tests that check the contents of 'other_dictionary'.
```

**Question 6: Data Science Modules (1 point)**

This question is just about importing the more core data science modules. We won't start using them yet. 

This is just to test that you have them available and working (this is the only time we'll ask you to write out your own import modules in assignments). This should work on datahub automatically, but we want to be sure *and* want to be sure you're familiar with how to do this.  

If you are unable to import them, ask for help in office hours.

Import the `numpy`, `scipy`, `pandas`, `matplotlib.pyplot`, `sklearn` modules as `np`, `sp`, `pd`, `plt`, `skl`, respectively.

<!-- grade_id: imports, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
import scipy as sp
import pandas as pd
import matplotlib.pyplot as plt
import sklearn as skl
```

<!-- grade_id: autograder-tests, points: 1, grade: true, solution: false -->
```python
# Tests for Q6
assert(np.array)
assert(sp.arcsinh)
assert(pd.DataFrame)
assert(plt.plot)
assert(skl.set_config)
```

## Done! Re-start & run all cells to be sure that everything passes, validate, and submit on DataHub!